evaluation_interval: 1m

rule_files:
- remotesecret_availability_recording_rules.yaml

tests:
- interval: 1m
  name: RSExporterTest
  input_series:
    - series: "redhat_appstudio_remotesecret_secretstorage_system_available{service='remote-secret-controller-manager-metrics-service'}"
      values: "1 1 1 1 1"
    - series: "redhat_appstudio_remotesecret_secretstorage_system_available{service='remote-secret-controller-manager-metrics-service1'}"
      values: "0 0 0 0 0"
    - series: "redhat_appstudio_remotesecret_secretstorage_system_available{service='remote-secret-controller-manager-metrics-service2'}"
      values: "0 1 0 1 0"
  promql_expr_test:
    - expr: konflux_up
      eval_time: 5m
      exp_samples:
        - labels: konflux_up{service='remote-secret-controller-manager-metrics-service', check='secretstorage'}
          value: 1
        - labels: konflux_up{service='remote-secret-controller-manager-metrics-service1', check='secretstorage'}
          value: 0
        - labels: konflux_up{service='remote-secret-controller-manager-metrics-service2', check='secretstorage'}
          value: 0
