evaluation_interval: 1m

rule_files:
  - prometheus.availability_metric_grafana_alerts.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="grafana", check="prometheus-appstudio-ds"}'
        values: '1x15'

    alert_rule_test:
      - eval_time: 15m
        alertname: KonfluxUpGrafanaAlert

  - interval: 1m
    input_series:
      - series: 'konflux_up{service="has", check="github"}'
        values: '_x15'

    alert_rule_test:

      - eval_time: 15m
        alertname: KonfluxUpGrafanaAlert
        exp_alerts:
          - exp_labels:
              severity: warning
              check: prometheus-appstudio-ds
              service: grafana
            exp_annotations:
              summary: >-
                Availability metric 'konflux_up' is missing.
              description: >-
                Availability metric 'konflux_up' is missing for check prometheus-appstudio-ds
                on service grafana for more than 5 minutes.
              alert_routing_key: 'appstudio-grafana'
