evaluation_interval: 1m

rule_files:
  - prometheus.availability_metric_remotesecret_alerts.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'konflux_up{service="remote-secret-controller-manager-metrics-service", check="secretstorage"}'
        values: '1x15'

    alert_rule_test:
      - eval_time: 15m
        alertname: KonfluxUpRemoteSecretAlert

  - interval: 1m
    input_series:
      - series: 'konflux_up{service="has", check="github"}'
        values: '_x15'

    alert_rule_test:
      - eval_time: 15m
        alertname: KonfluxUpRemoteSecretAlert
        exp_alerts:
          - exp_labels:
              severity: warning
              check: secretstorage
              service: remote-secret-controller-manager-metrics-service
            exp_annotations:
              summary: >-
                Availability metric 'konflux_up' is missing.
              description: >-
                Availability metric 'konflux_up' is missing for check secretstorage
                on service remote-secret-controller-manager-metrics-service for
                more than 5 minutes.
              alert_routing_key: 'appstudio-grafana'
